rule_files:
  - alerts.yaml
  - rules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: 'thanos_sidecar_last_heartbeat_success_time_seconds{namespace="production", job="thanos-sidecar", instance="thanos-sidecar-0"}'
    values: '5 10 43 17 11 _x5 0x10'
  - series: 'thanos_sidecar_last_heartbeat_success_time_seconds{namespace="production", job="thanos-sidecar", instance="thanos-sidecar-1"}'
    values: '4 9 42 15 10 _x5 0x10'
  promql_expr_test:
    - expr: time()
      eval_time: 1m
      exp_samples:
        - labels: '{}'
          value: 60
    - expr: time()
      eval_time: 2m
      exp_samples:
        - labels: '{}'
          value: 120
    - expr: time() - max(timestamp(thanos_sidecar_last_heartbeat_success_time_seconds{job="thanos-sidecar"})) by (job, instance)
      eval_time: 5m
      exp_samples:
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-0"}'
          value: 60
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-1"}'
          value: 60
    - expr: time() - max(timestamp(thanos_sidecar_last_heartbeat_success_time_seconds{job="thanos-sidecar"})) by (job, instance)
      eval_time: 6m
      exp_samples:
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-0"}'
          value: 120
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-1"}'
          value: 120
    - expr: time() - max(timestamp(thanos_sidecar_last_heartbeat_success_time_seconds{job="thanos-sidecar"})) by (job, instance)
      eval_time: 7m
      exp_samples:
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-0"}'
          value: 180
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-1"}'
          value: 180
    - expr: time() - max(timestamp(thanos_sidecar_last_heartbeat_success_time_seconds{job="thanos-sidecar"})) by (job, instance)
      eval_time: 8m
      exp_samples:
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-0"}'
          value: 240
        - labels: '{job="thanos-sidecar", instance="thanos-sidecar-1"}'
          value: 240
  alert_rule_test:
<<<<<<< HEAD
    - eval_time: 1m
      alertname: ThanosSidecarUnhealthy
    - eval_time: 2m
      alertname: ThanosSidecarUnhealthy
    - eval_time: 3m
      alertname: ThanosSidecarUnhealthy
    - eval_time: 5m
      alertname: ThanosSidecarUnhealthy
    - eval_time: 8m
      alertname: ThanosSidecarUnhealthy
      exp_alerts:
      - exp_labels:
          severity: critical
          job: thanos-sidecar
          instance: thanos-sidecar-0
        exp_annotations:
          description: 'Thanos Sidecar thanos-sidecar thanos-sidecar-0 is unhealthy for more than 240 seconds.'
          runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
          summary: 'Thanos Sidecar is unhealthy.'
      - exp_labels:
          severity: critical
          job: thanos-sidecar
          instance: thanos-sidecar-1
        exp_annotations:
          description: 'Thanos Sidecar thanos-sidecar thanos-sidecar-1 is unhealthy for more than 240 seconds.'
          runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
          summary: 'Thanos Sidecar is unhealthy.'
    - eval_time: 10m
      alertname: ThanosSidecarUnhealthy
=======
  - eval_time: 1m
    alertname: ThanosSidecarUnhealthy
  - eval_time: 2m
    alertname: ThanosSidecarUnhealthy
  - eval_time: 3m
    alertname: ThanosSidecarUnhealthy
  - eval_time: 10m
    alertname: ThanosSidecarUnhealthy
    exp_alerts:
    - exp_labels:
        severity: critical
        job: thanos-sidecar
        instance: thanos-sidecar-0
        namespace: production
      exp_annotations:
        description: 'Thanos Sidecar thanos-sidecar-0 in production is unhealthy for 600 seconds.'
        runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
        summary: 'Thanos Sidecar is unhealthy.'
    - exp_labels:
        severity: critical
        job: thanos-sidecar
        instance: thanos-sidecar-1
        namespace: production
      exp_annotations:
        description: 'Thanos Sidecar thanos-sidecar-1 in production is unhealthy for 600 seconds.'
        runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
        summary: 'Thanos Sidecar is unhealthy.'
  - eval_time: 11m
    alertname: ThanosSidecarUnhealthy
    exp_alerts:
    - exp_labels:
        severity: critical
        job: thanos-sidecar
        instance: thanos-sidecar-0
        namespace: production
      exp_annotations:
        description: 'Thanos Sidecar thanos-sidecar-0 in production is unhealthy for 660 seconds.'
        runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
        summary: 'Thanos Sidecar is unhealthy.'
    - exp_labels:
        severity: critical
        job: thanos-sidecar
        instance: thanos-sidecar-1
        namespace: production
      exp_annotations:
        description: 'Thanos Sidecar thanos-sidecar-1 in production is unhealthy for 660 seconds.'
        runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
        summary: 'Thanos Sidecar is unhealthy.'
  - eval_time: 12m
    alertname: ThanosSidecarUnhealthy
    exp_alerts:
    - exp_labels:
        severity: critical
        job: thanos-sidecar
        instance: thanos-sidecar-0
        namespace: production
      exp_annotations:
        description: 'Thanos Sidecar thanos-sidecar-0 in production is unhealthy for 720 seconds.'
        runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
        summary: 'Thanos Sidecar is unhealthy.'
    - exp_labels:
        severity: critical
        job: thanos-sidecar
        instance: thanos-sidecar-1
        namespace: production
      exp_annotations:
        description: 'Thanos Sidecar thanos-sidecar-1 in production is unhealthy for 720 seconds.'
        runbook_url: 'https://github.com/thanos-io/thanos/tree/main/mixin/runbook.md#alert-name-thanossidecarunhealthy'
        summary: 'Thanos Sidecar is unhealthy.'
>>>>>>> 03e314dd (Add namespaces to alerts)
